How Absence Seizures Impair Sensory Perception: Insights from Awake fMRI and Simulation Studies in Rats

Author response:

The following is the authors’ response to the previous reviews.

eLife assessment

This valuable work performed fMRI experiments in a rodent model of absence seizures. The results provide new information regarding the brain's responsiveness to environmental stimuli during absence seizures. The authors suggest reduced responsiveness occurs during this type of seizure, and the evidence leading to the conclusion is solid, although reviewers had divergent opinions.

Public Reviews:

Reviewer #1 (Public Review):

In this paper, the effects of two sensory stimuli (visual and somatosensory) on fMRI responsiveness during absence seizures were investigated in GEARS rats with concurrent EEG recordings. SPM analysis of fMRI showed a significant reduction in whole-brain responsiveness during the ictal period compared to the interictal period under both stimuli, and this phenomenon was replicated in a structurally constrained whole-brain computational model of rat brains.

The conclusion of this paper is that whole-brain responsiveness to both sensory stimuli is inhibited and spatially impeded during seizures.

Reviewer #2 (Public Review):

Summary:

This study examined the possible affect of spike-wave discharges (SWDs) on the response to visual or somatosensory stimulation using fMRI and EEG. This is a significant topic because SWDs often are called seizures and because there is non-responsiveness at this time, it would be logical that responses to sensory stimulation are reduced. On the other hand, in rodents with SWDs, sensory stimulation (a noise, for example) often terminates the SWD/seizure.

In humans, these periods of SWDs are due to thalamocortical oscillations. A certain percentage of the normal population can have SWDs in response to photic stimulation at specific frequencies. Other individuals develop SWDs without stimulation. They disrupt consciousness. Individuals have an absent look, or "absence", which is called absence epilepsy.

The authors use a rat model to study the responses to stimulation of the visual or somatosensory systems during and in between SWDs. They report that the response to stimulation is reduced during the SWDs. While some data show this nicely, the authors also report on lines 396-8 "When comparing statistical responses between both states, significant changes (p<0.05, cluster-) were noticed in somatosensory auditory frontal..., with these regions being less activated in interictal state (see also Figure 4). That statement is at odds with their conclusion. I do not see that this issue was addressed.

See comments below starting with “We acknowledge the reviewer…”.

They also conclude that stimulation slows the pathways activated by the stimulus. I do not see any data proving this. It would require repeated assessments of the pathways in time. This issue was not addressed.

See comments below starting with “We acknowledge the reviewer…”.

The authors also study the hemodynamic response function (HRF) and it is not clear what conclusions can be made from the data. This is still an issue. No conclusions appear to be possible to make.

See comments below starting with “We acknowledge the reviewer…”.

Finally, the authors use a model to analyze the data. This model is novel and while that is a strength, its validation is unclear. The authors did not add any validation of their model.

See comments below starting with “We acknowledge the reviewer…”.

Strengths:

Use of fMRI and EEG to study SWDs in rats.

Weaknesses:

Several aspects of the Methods and Results were improved but some are still are unclear.

We acknowledge the reviewer for the concerns of we not addressing the comments above. However, we emphasize that most of the comments were addressed in the already sent “Response to Review Comments” and in the updated manuscript. Here we repeat the responses and provide also additional clarifications to some of the comments.

We thank the reviewer for noting the discrepancy in the statement of “less activated in interictal state”. The statement should have been written vice versa. We also address that the direction of activation change between groups can be misinterpreted based on statistical maps itself (Figure 3) where only statistical changes are visible and not the polarity of response (can be seen in Figure 4). Therefore, we have made a following changes in the section 3.3.: “There were more voxels with significant changes of activity during interictal state compared to ictal state (136% more). Comparing the statistical responses between interictal and ictal states revealed significant changes (p<0.05, cluster-level corrected) in the visual, somatosensory, and medial frontal cortices. In the ictal state, these regions showed significant hemodynamic decreases when comparing to interictal state, and these polarity changes can be seen the hemodynamic response functions (Figure 4).”

We agree with the reviewer that there are no data showing slowing of the pathways in response to stimulus. However, we are a bit confused about this comment, as to what part in conclusion section it refers to. We did not intentionally claim that stimulation slows the activated pathways in the manuscript.

Reviewer is right that strong claims cannot be made from HRF by itself. Therefore, we have avoided to such phrasing throughout the manuscript. In the conclusion section, we speculate that HRF decreases “could play a role in decreased sensory perception” but also state that “further studies are required”. The observed HRF decreases (rather than increases) in the cortex when stimulation was applied during SWD, was discussed in section 4.4., where we speculated that neuronal suppression (possible apparent in negative HRFs) caused by SWD can prevent responsiveness. Conclusion now states the following: “Moreover, the detected decreases in the cortical HRF when sensory stimulation was applied during spike-and-wave discharges, could play a role in decreased sensory perception. Further studies are required to evaluate whether this HRF change is a cause or a consequence of the reduced neuronal response.”

We point out that the main validation of the model and its details were provided in the previous answer to the reviewer and added to the manuscript. The model presented in the paper is based on a mean-field formalism that captures neuronal activity at the mesoscale level. This mean-field formalism is derived via a detailed statistical description of the activity of a spiking neuronal population of excitatory and inhibitory with conductance-based synaptic interactions. Thus, the validation of the mean-field model is performed via direct comparison between the dynamics obtained from the mean-field model and the dynamics obtained from the underlying spiking neural network model. This comparison is shown in the supplementary material of the manuscript, where the transition studied in the paper between interictal (asynchronous irregular activity) and ictal (SWD dynamics) activity, which is predicted by the mean-field model, is indeed observed in the underlying spiking neuronal model. The existence of these two types of dynamics and the transition between them is the main component of the model used to build the analysis of the responsiveness performed in the paper (which has been properly validated).

Reviewer #3 (Public Review):

Summary:

This is an interesting paper investigating fMRI changes during sensory (visual, tactile) stimulation and absence seizures in the GAERS model. The results are potentially important for the field and do suggest that sensory stimulation may not activate brain regions normally during absence seizures. But the findings are limited by substantial methodological issues that do not enable fMRI signals related to absence seizures to be fully disentangled from fMRI signals related to the sensory stimuli.

Strengths:

Investigating fMRI brain responses to sensory stimuli during absence seizures in an animal model is a novel approach with potential to yield important insights.

Use of an awake, habituated model is a valid and potentially powerful approach.

Weaknesses:

The major difficulty with interpreting the results of this study is that the duration of the visual and tactile stimuli were 6 seconds, which is very close to the mean seizure duration per Table 1. Therefore the HRF model looking at fMRI responses to visual or auditory stimuli occurring during seizures was simultaneously weighting both seizure activity and the sensory (visual or auditory) stimuli over the same time intervals on average. The resulting maps and time courses claiming to show fMRI changes from visual or auditory stimulation during seizures will therefore in reality contain some mix of both sensory stimulation-related signals and seizure-related signals. The main claim that the sensory stimuli do not elicit the same activations during seizures as they do in the interictal period may still be true. But the attempts to localize these differences in space or time will be contaminated by the seizure related signals.

In their response to this comment the authors state that some seizures had longer than average duration, and that they attempted to model the effects of both seizures and sensory stimulation. However these factors do not mitigate the concern because the mean duration of seizures and sensory stimulation remain nearly identical, and the models used therefore will not be able to effectively separate signals related to seizures and related to sensory stimulation.

Regressors for seizures were formed by including periods of seizures without any stimulation present. In theory, if seizures were perfectly modeled by the regressor, the left variance is completely orthogonal to the main effect of the stimulus. Furthermore, only the cases where the seizures are longer than the stimulus are used to calculate the responsiveness of the stimulus (while the cases where the seizures are shorter than the stimulus are used as nuisance regressors to account for error variance). However, we agree with the reviewer that in practice all effects of the seizure cannot be removed completely from the effect of stimulus. We have addressed this concern in the “physiologic and methodology consideration” section: “We note a caution that presented maps and time courses showing fMRI changes from visual or whisker stimulation during seizures may contain a mixture of both sensory stimulation-related signals and seizure-related signals. To minimize this contamination in the linear model used, we considered both stimulation and seizure-only states as regressors of interest and used seizure-only responses as nuisance regressors to account for error variance. Thereby, the effects caused by the stimulation should be separated as much as possible from the effects caused by the seizure itself.”

The claims that differences were observed for example between visual cortex and superior colliculus signals with visual stim during seizures vs interictal remain unconvincing due to above.

Maps shown in Figure 3 do not show clear changes in the areas claimed to be involved.

In their response the authors enlarged the cross sections. However there are still discrepancies between the images and the way they are described in the text. For example, in the Results text the authors say that comparing the interictal and ictal states revealed less activation in the somatosensory cortex during the ictal than during the interictal state, yet Figure 3 bottom row left shows greater activation in somatosensory cortex in this contrast.

We note that the direction of activation change between groups can be misinterpreted based on statistical maps itself (Figure 3) where only statistical changes are visible and not the polarity of response (can be seen in Figure 4). Therefore, we have made the following changes to the section 3.3.: “There were more voxels with significant changes of activity during interictal state compared to ictal state (136% more). Comparing the statistical responses between interictal and ictal states revealed significant changes (p<0.05, cluster-level corrected) in the visual, somatosensory, and medial frontal cortices. In the ictal state, these regions showed significant hemodynamic decreases when comparing to interictal state, and these polarity changes can be seen the hemodynamic response functions (Figure 4).”

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

Authors have revised this paper with a lot of detail. The paper can be accepted for publication in this version.

Reviewer #2 (Recommendations For The Authors):

Reviewer #1

(1) The analysis in this paper does not directly answer the scientific question posed by the authors, which is to explore the mechanisms of the reduced brain responsiveness to external stimuli during absence seizures (in terms of altered information processing), but merely characterizes the spatial involvement of such reduced responsiveness. The same holds for the use of mean-field modeling, which merely reproduces experimental results without explaining them mechanistically as what the authors have claimed at the head of the paper.

We agree with the reviewer that the manuscript does not answer specifically about the mechanisms of reduced brain responsiveness. The main scientific question addressed in the manuscript was to compare whole-brain responsiveness of stimulus between ictal and interictal states. The sentence that can lead to misinterpretations in the manuscript abstract: "The mechanism underlying the reduced responsiveness to external stimulus remains unknown." was therefore modified to the following "The whole-brain spatial and temporal characteristics of reduced responsiveness to external stimulus remains unknown".

This change did not address the issue. The problem is that there is no experimentation to address the underlying mechanisms of the results. I also think the changed language in the abstract is less clear than the original.

We fully agree that this manuscript does not answer or claim to be answering about the mechanisms of reduced brain responsiveness. The main scientific question addressed in the manuscript was to compare whole-brain responsiveness of stimulus between ictal and interictal states, by means of hemodynamics and mean-field simulation.

We have changed the language of the abstract to the following:

“In patients suffering absence epilepsy, recurring seizures can significantly decrease their quality of life and lead to yet untreatable comorbidities. Absence seizures are characterized by spike-and-wave discharges on the electroencephalogram associated with a transient alteration of consciousness. However, it is still unknown how the brain responds to external stimuli during and outside of seizures.

This study aimed to investigate responsiveness to visual and somatosensory stimulation in GAERS, a well-established rat model for absence epilepsy. Animals were maintained in a non-curarized awake state allowing for naturally occurring seizures to be produced inside the magnet. They were imaged continuously using a quiet zero-echo-time functional magnetic resonance imaging (fMRI) sequence. Sensory stimulations were applied during interictal and ictal periods. Whole brain responsiveness and hemodynamic responses were compared between these two states. Additionally, a mean-field simulation model was used to mechanistically explain the changes of neural responsiveness to visual stimulation between interictal and ictal states.

Results showed that, during a seizure, whole-brain responses to both sensory stimulations were suppressed and spatially hindered. In several cortical regions, hemodynamic responses were negatively polarized during seizures, despite the application of a stimulus. The simulation experiments also showed restricted propagation of spontaneous activity due to stimulation and so agreed well with fMRI findings. These results suggest that sensory processing observed during an interictal state is hindered or even suppressed by the occurrence of an absence seizure, potentially contributing to decreased responsiveness during this absence epileptic process.”

The authors also study the hemodynamic response function (HRF) and it is not clear what conclusions can be made from the data.

The response of the authors did not clarify this issue. Instead, they explained why they examined HRF and that they can only speculate what the data means.

Reviewer is right that strong claims cannot be made from HRF by itself. Therefore, we have avoided to such phrasing throughout the manuscript. In the conclusion section, we speculate that HRF decreases “could play a role in decreased sensory perception” but also state that “further studies are required”.

Finally, the authors use a model to analyze the data. This model is novel and while that is a strength, its validation is unclear. The conclusion is that the modeling supports the conclusions of the study, which is useful.

Details about the model were added.

This is not entirely satisfactory because there is still no validation of the model.

We point out that the main validation of the model and its details were provided in the previous answer to the reviewer and added to the manuscript. The model presented in the paper is based on a mean-field formalism that captures neuronal activity at the mesoscale level. This mean-field formalism is derived via a detailed statistical description of the activity of a spiking neuronal population of excitatory and inhibitory with conductance-based synaptic interactions. Thus, the validation of the mean-field model is performed via direct comparison between the dynamics obtained from the mean-field model and the dynamics obtained from the underlying spiking neural network model. This comparison is shown in the supplementary material of the manuscript, where the transition studied in the paper between interictal (asynchronous irregular activity) and ictal (SWD dynamics) activity, which is predicted by the mean-field model, is indeed observed in the underlying spiking neuronal model. The existence of these two types of dynamics and the transition between them is the main component of the model used to build the analysis of the responsiveness performed in the paper (which has been properly validated).

How is ROI defined in this paper? What type of atlas is used?

Anatomical ROIs were drawn based on Paxinos and Watson rat brain atlas 7th edition. Region was selected if there were statistically significant activations detected inside that region, based on activation maps. We clarified the definition of ROI as the following:

"Anatomical ROIs, based on Paxinos atlas (Paxinos and Watson rat brain atlas 7th edition), were drawn on the brain areas where statistical differences were seen in activation maps."

This is helpful, but the unstained brain does not show the borders of the areas. Therefore just saying an atlas was used is not enough. How in an unstained brain can the areas be accurately outlined?

Areas of the brain were differentiated by co-registering the functional MRI images with an T1-weighted anatomical reference brain that was created on site from the same data set that was used for the manuscript. Potential co-registration inaccuracies created by using a reference brain measured in different site, sequence and a rat strain can be thus avoided. T1-images create sufficient contrast to differentiate main brain areas, but for more accurate border definition (e.g., to differentiate different thalamic nuclei), a coordinate system of the atlas and coordinates known in the used anatomical brain, were used to pinpoint exact borders of the brain areas.

Reviewer #2

The following also is not precise:

"Although seizures are initially triggered by hyperactive somatosensory cortical neurons, the majority of neuronal populations are deactivated rather than activated during the seizure, resulting in an overall decrease in neuronal activity during SWD (McCafferty et al. 2023)."

What neuronal populations? Cortex? Which neurons in the cortex? Those projecting to the thalamus? What about thalamocortical relay cells? Thalamic gabaergic neurons?

Please check that these issues were corrected.

The issues were addressed as follows:

“Although SWDs are initially triggered by hyperactive somatosensory cortical neurons, neuronal firing rates, especially in majority of frontoparietal cortical and thalamocortical relay neurons, are decreased rather than increased during SWD, resulting in an overall decrease in activity in these neuronal populations (McCafferty et al., 2023). Previous fMRI studies have demonstrated blood volume or BOLD signal decreases in several cortical regions including parietal and occipital cortex, but also, quite surprisingly, increases in subcortical regions such as thalamus, medulla and pons (David et al., 2008; McCafferty et al., 2023).”

Results

After removing problematic animals and sessions, was there sufficient power? There probably wasn't enough to determine sex differences.

After removing problematic sessions, we found statistically significant results (multiple comparison corrected) results in both activation maps, and hemodynamic responses. To determine sex differences, there were not enough animals for statistical findings (p>0.05).

This is not the question. The question is whether there was sufficient power.

A simple power calculation was performed as follows: considering a t-test, a risk alpha of 0.05, a power of 0.8, matched pairs (seizure/control), we can detect an effect size of 0.37 with our 4 animals, considering repeated measurements (4 sessions/animal x 11 seizure/control pairs per session). This is now mentioned in the manuscript.

Table 1 has no statistical comparisons.

Table 1 is purely an illustration of stimulation and seizure occurrence. There is no specific interest to compare stimulation types (in what state of seizure it occurred) as it does not provide any meaningful inferences to the study.

Table 1 could be improved by statistics. More could be said and there would be justification to include it.

We thank the reviewer for the suggestion, but as it is yet unclear to what statistical comparison would be feasible to do, we opt to leave it out.

Statistical activation maps - it is not clear how this was done.

Creation of statistical maps are explained in section 2.5.3.

This section is not clear.

We have added a reference (https://doi.org/10.1002/hbm.460020402) for readers to familiarize themselves with the concept of statistical parametric mapping.

Fig 3 "F-contrast maps." Please explain.

Creation of statistical maps are explained in section 2.5.3.

This section is unclear.

We have added a reference (https://doi.org/10.1002/hbm.460020402) for readers to familiarize themself with the concept of statistical parametric mapping.

Reviewer #3 (Recommendations For The Authors):

Aside from the concerns listed as weaknesses above which were not addressed, most of the more minor comments were addressed by the authors in the resubmission. However, the comment below was not addressed because it is impossible to see any firing rate changes elicited by sensory stimuli (if they are present) due to the scale during seizures. The seizure signals should be removed or accounted for by the model so that any possible sensory stimulus-related signals could be seen, and displayed on the same scale as firing rates without seizures. Prior comment (unaddressed) is repeated below:

Figure 6-figure supplement 1, the scales are very different for many of the plots so they are hard to compare. Especially in the ictal periods (D, E, F) it is hard to see if any changes are happening during ictal stimulation similar to interictal stimulation due to very different scales. The activity related to SWD is so large that it overshadows the rest, and perhaps should be subtracted out.

These two comments were addressed and replied in the previous round of reviews. Regarding the different scales of the plots from Figure 6-figure supplement 1, we point out that all the plots in the same scale are already presented in Figure 6 of the main-text. Regarding the activity related to SWD and sensory stimulation, we remark that the effect of the stimulation should be (and was) evaluated with respect to the ongoing activity. All the results concerning the neuronal responsiveness presented in the paper evaluate the statistical significance of the changes in activity produced by the stimulation with respect to the ongoing activity (during ictal and interictal states respectively). For this reason, all the plots containing the time series of neuronal activity in the simulations include the ongoing activity (with SWD dynamics when present) for proper comparison and relevant analysis.

Additional changes:

In the section 3.2., the sentence: “In addition, responses were observed in the somatosensory cortex during a seizure state.” was removed for clarification purposes as deactivation rather than activation was observed in this brain area during a seizure state.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Absence seizures are characterized by regular and generalized spike-and-wave electrical patterns in the brain, resulting in unresponsiveness to environmental stimuli. In patients suffering absence epilepsy, recurring seizures can significantly decrease their quality of life and lead to yet untreatable comorbidities. The whole-brain spatial and temporal characteristics of reduced responsiveness to external stimulus remains unknown.

Revised version: In patients suffering absence epilepsy, recurring seizures can significantly decrease their quality of life and lead to yet untreatable comorbidities. Absence seizures are characterized by spike-and-wave discharges on the electroencephalogram associated with a transient alteration of consciousness. However, it is still unknown how the brain responds to external stimuli during and outside of seizures.

Earlier version: This study aimed to investigate responsiveness to visual and somatosensory whisker stimulation in GAERS, a well-established rat model for absence epilepsy. Animals were imaged continuously using a quiet zero-echo-time (ZTE) functional magnetic resonance imaging (fMRI) sequence while in a non-curarized awake state, allowing for naturally occurring seizures to be produced inside the 9.4T magnet. Sensory stimulations were applied in 28 fMRI sessions during interictal and ictal periods, as assessed by concurrent electroencephalography (EEG) recordings, and whole brain responsiveness and hemodynamic responses were compared between these two states. Additionally, a mean-field simulation model was used to mechanistically explain the changes of neural responsiveness to visual stimulation between interictal and ictal states.

Revised version: This study aimed to investigate responsiveness to visual and somatosensory stimulation in GAERS, a well-established rat model for absence epilepsy. Animals were maintained in a non-curarized awake state allowing for naturally occurring seizures to be produced inside the magnet. They were imaged continuously using a quiet zero-echo-time functional magnetic resonance imaging (fMRI) sequence. Sensory stimulations were applied during interictal and ictal periods. Whole brain responsiveness and hemodynamic responses were compared between these two states. Additionally, a mean-field simulation model was used to mechanistically explain the changes of neural responsiveness to visual stimulation between interictal and ictal states.

Earlier version: Results showed that whole-brain responses to both sensory stimulations were suppressed and spatially hindered during a seizure. In several cortical regions, hemodynamic responses were negatively polarized during seizures, despite the application of a stimulus. The simulation results agreed well with fMRI findings, showing restricted propagation of spontaneous activity due to stimulation. These results suggest that sensory processing observed during an interictal state is hindered or even suppressed by the occurrence of an absence seizure potentially contributing to decreased responsiveness.

Revised version: Results showed that, during a seizure, whole-brain responses to both sensory stimulations were suppressed and spatially hindered. In several cortical regions, hemodynamic responses were negatively polarized during seizures, despite the application of a stimulus. The simulation experiments also showed restricted propagation of spontaneous activity due to stimulation and so agreed well with fMRI findings. These results suggest that sensory processing observed during an interictal state is hindered or even suppressed by the occurrence of an absence seizure, potentially contributing to decreased responsiveness during this absence epileptic process.

Introduction

Absence seizures typically appear in children between the ages of 5 and 7, and are characterized by a sudden, brief impairment of consciousness, an interruption from ongoing activities, and an unresponsiveness to environmental stimuli (Loiseau et al., 1995). These seizures typically last for a few seconds to several minutes and include a regular electroencephalographic pattern known as spike-and-wave discharge (SWD) with a frequency of 2-5 Hz (Szaflarski et al., 2010). In non-treated patients, absence seizures can occur from a few to hundreds of times per day (Loiseau et al., 1995). The impaired consciousness during absence seizures can be highly disabling. Indeed, absence epilepsy is not a benign condition and is often accompanied by severe neuropsychiatric comorbidities including impairment of attention, memory and mood (Crunelli et al., 2020). To better understand how these co-morbidities arise, it is important to investigate how information processing is altered between the ictal and interictal periods.

In absence epilepsy, human neuroimaging studies suggest that during a seizure, there may be a lack of conscious information processing due to impaired fronto-parietal network, arousal systems in the thalamus and brainstem (Blumenfeld, 2012), or default mode network (Luo et al., 2011). Likewise, multiple human studies on absence seizure have demonstrated lack of responsiveness to external stimuli such as commands and questions, potentially caused by focal disruption of information processing in specific corticothalamic networks (Blumenfeld, 2012). However, to our knowledge, current neuroimaging studies on absence epilepsy have been conducted in a resting state, without external stimulus. Therefore, it remains unclear how the brain manages environmental stimuli during absence discharges, and more research is needed to understand how the brain responsiveness is affected during altered brain states.

The Genetic Absence Epilepsy Rats from Strasbourg (GAERS) is a well-established model of absence epilepsy. It is based on the selection of rats exhibiting spontaneous SWD and recapitulates most electrophysiological, behavioral and pharmacological features of human absence epilepsy (Depaulis et al., 2016). Although the origin of absence seizures is not fully understood, current studies on rat models of absence seizures suggest that they arise from excitatory drive in the barrel field of the somatosensory cortex (David et al., 2008; Meeren et al., 2002; Polack et al., 2007) and then propagate to other structures (David et al., 2008) including thalamus, knowing to play an essential role during the ictal state (Huguenard, 2019). Notably, thalamic subnetwork is believed to play a role in coordinating and spacing of SWDs via feedforward inhibition together with burst firing patterns. This leads to the rhythms of neuronal silence and activation periods that are detected in SWD waves and spikes (Huguenard, 2019; McCafferty et al., 2023).

Previous results on GAERS have indicated that, during an absence seizure, hyperactive electrophysiological activity in the somatosensory cortex can contribute to bilateral and regular SWD firing patterns in most parts of the cortex. These patterns propagate to different cortical areas (retrosplenial, visual, motor and secondary sensory), basal ganglia, cerebellum, substantia nigra and thalamus (David et al., 2008; Polack et al., 2007). Although SWDs are initially triggered by hyperactive somatosensory cortical neurons, neuronal firing rates, especially in majority of frontoparietal cortical and thalamocortical relay neurons, are decreased rather than increased during SWD, resulting in an overall decrease in activity in these neuronal populations (McCafferty et al., 2023). Previous fMRI studies have demonstrated blood volume or BOLD signal decreases in several cortical regions including parietal and occipital cortex, but also, quite surprisingly, increases in subcortical regions such as thalamus, medulla and pons (David et al., 2008; McCafferty et al., 2023). In line with these findings, graph-based analyses have shown an increased segregation of cortical networks from the rest of the brain (Wachsmuth et al., 2021). Altogether, alterations in these focal networks in the animal models of epilepsy impairs cognitive capabilities needed to process specific concurrent stimuli during SWD and therefore could contribute to the lack of behavioral responsiveness (Chipaux et al., 2013; Luo et al., 2011; Meeren et al., 2002; Studer et al., 2019), although partial voluntary control in certain stimulation schemes can be still present (Taylor et al., 2017).

The objective of this study was to investigate changes in whole brain responsiveness to sensory stimuli during ictal and interictal states, using the GAERS animal model. To avoid the potential impact of anesthetic agents on fMRI recordings and because absence seizures can be only observed in awake animals, EEG-fMRI was performed in GAERS trained to remain still and awake, i.e. non-curarized and non-anesthetized, using a previously validated protocol (Paasonen et al., 2018; Stenroos et al., 2018). For this purpose, the functionality of ZTE sequence was first piloted, and selected over traditional EPI sequence for its lower acoustic noise and reduced magnetic susceptibility artefacts. The selected MRI sequence thus appeared optimal for awake EEG-fMRI measurements. Visual and somatosensory whisker stimulations were used during both ictal and interictal states, as previous research has shown altered electrophysiological activity and behavior in these sensory systems during seizures (Meeren et al., 2002; Meyer et al., 2018; Pavone et al., 2001; Studer et al., 2019). By using both stimulation schemes, we aimed to investigate alterations in brain responsiveness in each sensory system and to identify any common changes. To further describe the change when switching between ictal and interictal states, the whole-brain response to visual stimulus in each state was simulated, using a recent mean-field model (Di Volo et al., 2019).

Methods

Animals

All experiments were approved by the local animal welfare committee (Comité Local GIN, C2EA-04) and complied with EU guidelines (Directive 2010/63/EU). Every precaution was taken to minimize the number of animals used and stress to animals during experiments. A total of 11 adult 8–12-month age old GAERS rats were used (260 ± 21 g, 6 males, 5 females). After implantation of an EEG lead (see below), rats were individually housed in their cages. Animals were maintained on a 12/12 h light-dark cycle at room temperature of 22 ± 2°C, humidity of 50–60%. Food (Extrudat, vitamin-fortified, irradiated > 25 kGy) and water were available ad libitum.

EEG implantation

Carbon fiber (WPI Sarasota FL) electrodes were prepared as follows. Fiber leads were cut to small (∼30 mm in length) parts and the end of the electrode was exposed from insulating cover to leave a brush-like, ∼5 mm-length, tip. Rats were anesthetized with isoflurane (induction with 5%, a percentage gradually decreased to 1-2% for the surgery maintenance). The head was then shaved to remove the fur above the scalp, and the rat was positioned in a stereotaxic frame (David Kopf Instruments, Germany). After local lidocaine chlorhydrate injections (2g/100ml, 0,05 ml/site), the skull was exposed, cleaned with sterile 0.9% saline and hydrogen peroxide, and allowed to dry. Small holes were drilled halfway through the skull and carbon fiber electrodes were laid and glued with cyanoacrylate over the right motor cortex (AP: +2, ML: + 2,5 mm) and right primary somatosensory cortex (AP: −2,5, ML: + 3 mm) for seizure detection. An electrode working both as a reference and ground was placed on top of the cerebellum (AP: −12, ML: + 2 mm). The other end of electrode leads was inserted inside a plastic tube), which was positioned at the center of the skull. A thin layer of cyanoacrylate was applied to cover the skull. On top of the glue, a layer of dental cement (Selectaplus, DeguDent GmbH, Germany) was applied to finish the implant. Following the surgery, rats were individually caged to recover, and the welfare of the animals was closely monitored.

Animal habituation for awake imaging

A low noise zero echo time (ZTE) MRI sequence (see Wiesinger & Ho, 2022) was used with awake rats aiming for low-stress and low-motion functional imaging sessions that allows rats to produce spontaneous seizures. For the habituation, peak ZTE scanner noise was measured with an omnidirectional condenser microphone (MT830R, Audio-Technica Limited, Leeds, UK) and Audacity software (version 2.3.0, https://www.audacityteam.org/), similar to (Paasonen et al., 2020) and reproduced at equivalent sound pressure through a loudspeaker.

To habituate the rats for the fMRI experiments, a procedure based on a previous study was followed (Stenroos et al., 2018). Restraint parts compatible with standard Bruker rat bed and suitable for stimulation leads were designed using 123D CAD software (Autodesk, San Rafael, CA) and 3D-printed with Ultimaker 2 (Utrecht, Netherlands) using Acrylonitrile Butadiene Styrene (ABS) plastic (Figure 1). The rats’ bodies were restrained with a soft and elastic foamed plastic, hind legs were taped together, and front legs were taped loosely together to the side of the body. Front teeth were secured with a carbon fiber teeth bar, head with a nose cone, and neck and shoulders with a neck and shoulder bars, respectively. Silicone ear plugs were used to minimize experienced noise. Rats were habituated to restraint and MRI ZTE gradient noises by gradually increasing session times from 15 minutes to 60 minutes per day for 8 days before the first fMRI experiments. The length of habituation period was selected based on pilot experiments to provide low-motion data therefore giving rats a better chance to be in a low-stress state and thus produce absence seizures inside the magnet. Pressure pillow and video camera were used to estimate physiological state, via breathing rate, and motion level, respectively. During the last habituation session, EEG was measured to confirm that the rats produced a sufficient amount of absence seizures (10 or more per session). Total of 3-5 fMRI experiments were conducted per rat within a 1–3-week period. In case rats were re-imaged more than 1 week after the preceding experiment, an additional 2-day habituation period was conducted (Figure 1–figure supplement 2). Before and after each habituation session, rats were given a treat of 1% sugar water and/or 3 chocolate cereals as positive reinforcement.

Sensory stimulation, EEG setup and example EEG traces.

For visual stimulation, optical fiber cables were positioned bilaterally close to the eyes (A1); For whisker stimulation, plastic tips guided air flow bilaterally to whiskers (A2). For EEG, carbon fiber leads were connected to electrodes coming from a plastic tube (A3). EEG traces are illustrated at 3 different temporal scale, during ictal and interictal states and with stimulation onset marks (B). Color tags mark stimulation onsets (purple) and the 6-second stimulation blocks during baseline (blue), seizure (red) and when stimulation ended a seizure (orange). The onset marks were added post-hoc, based on recorded TTL events.

EEG-MRI protocol

MRI acquisitions were carried out at 9.4 T (Biospec Avance III HD, Bruker, Ettlingen, Germany; IRMaGe facility) using Paravision 7. In-house made transmit-receive loop coil with a 22 mm inner diameter was designed to host the EEG lead and to be compatible with the rat restraining cradle (Figure 2). 3D-printed sledge was designed to cover the circuit board of the coil and partly stabilize coaxial cable and the loop. Before connecting carbon fiber electrode leads, the loop was placed around the implant and sledge was fastened to the animal cradle with masking tape. Next, electrode leads were connected to the other end of the EEG cable. The cable was carefully secured on top of 3D-holder, without touching the animal to avoid breathing and motion artifacts.

EEG-fMRI setup and illustration of MRI-data quality.

MRI transmit-receiver loop coil placed around the implant (A), spatial signal-to-noise ratios of an illustrative high resolution T1-FLASH (B), and low resolution ZTE image (C), temporal signal-to-noise ratios (tSNR) of ZTE data (D) from a one example animal and average framewise displacement (red) with the standard deviation (grey area) across all sessions included to analyses (E).

fMRI optimization

To select the optimal fMRI sequence for the study design, standard gradient-echo echo-planar imaging (GE-EPI) and zero echo time imaging (ZTE; see section 2.4.3 for details) were compared. Peak acoustic noise levels produced by the MRI scanner during the sequence run, level of spatial distortions on MRI images caused by electrodes, gradient switching artifacts on EEG, and functional contrast were evaluated (Figure 2, Figure 2–figure supplement 1). Motion incidences between this study and previous study using EPI sequence (Stenroos et al., 2018) were compared. ZTE was selected over EPI sequence as it was quieter, produced less susceptibility artifacts, had less noise on EEG recordings and led to less animal motion (see Results section). Low noise level of ZTE is particularly important factor for seizure appearance, as GAERS rats only experience seizures when they are awake and in a calm state. Eventually, the ZTE parameters were adapted to optimize the functional contrast to noise ratio together with the temporal and spatial resolutions, based on previous reports (MacKinnon et al. ISMRM 2021) and our own pilot acquisitions. Note that ZTE is sensitive to change in blood flow (Lehto et al., 2017) and not to BOLD contrast as in EPI, and therefore also more direct measure of neuronal activation. To assess the overall data quality and functionality of the transceiver loop coil, spatial and temporal signal to noise ratios (SNRs) were assessed in one representative rat (Figure 2). Spatial SNR was obtained as the ratio between the mean signal intensity in an area of interest (the cortex) and the standard deviation of the background signal, whereas temporal SNR was obtained voxel-wise by dividing the mean by the standard deviation of the normalized fMRI signal.

Animal installation

Before awake EEG-fMRI measurements, rats were anesthetized with isoflurane (5% induction, 2% maintenance in 30% O2/ 70% N2). Next, animals were wrapped and taped for awake imaging (section 2.3) and the animals were moved to the MRI scanner and restrained to the 3D-printed holder. The respiration rate was measured with a pressure pillow placed under animals and a Biopac amplifier system (Goleta, CA, USA). Temperature was not measured in awake condition to avoid causing any harms to the rectum due to motion. We relied on pilot calibration of the temperature of heated water circulating inside the animal bed to maintain the normal body temperature of ∼37 °C. Lastly, EEG leads were connected, and animals were pushed to the center of the magnet bore. After routine preparation steps in the MRI console were done, isoflurane was turned off. Once animals woke up from anesthesia, they were left to rest in a quiet bore for 5-15 min so that spontaneous seizures started to emerge, after which functional imaging was started.

EEG-MRI acquisition and stimulations parameters

Anatomical imaging was conducted with a T1-FLASH sequence (repetition time: 530 ms, echo time: 4 ms, flip angle 18°, bandwidth 39,682 kHz, matrix size 128 x 128, 51 slices, field-of-view 32 x 32 mm², spatial resolution 0.25 x 0.25 x 0.5 mm3). fMRI was performed with a 3D ZTE sequence (repetition time: 0.971 ms, echo time: 0 ms, flip angle 4°, pulse length 1 µs, bandwidth 150 kHz, oversampling 4, matrix size 60 x 60 x 60, field-of-view 30 x 30 x 60 mm3, spatial resolution of 0.5 x 0.5 x 1 mm3, polar under sampling factor 5.64, number of projections 2060 resulting to a volume acquisition time of about 2 s (look Wiesinger & Ho, 2022 for parameter explanations). A total of 1350 volumes (45 min) were acquired.

Visual (n=14 sessions, 5 rats) and somatosensory whisker (n=14 sessions, 4 rats) stimulations fMRI measurements were performed (Figure 1a). Stimulus duration, frequency and pulse lengths were automatically controlled by an Arduino chip (Arduino Uno Rev3). For the visual stimulation, light pulses (3 Hz, 6 s total length, pulse length 166 ms) were produced by a blue led (wavelength of 470 nm) and light was guided through two optical fibers to the front of the rat’s eyes.

For the somatosensory stimulation, an air-pressure valve (PMI-200 pressure micro-injector, Dagan corporation, Minneapolis, USA) controlled by the Arduino chip produced air flow pulses (2 Hz, 6 s total length, pulse length 250 ms) with a pressure of 2.3 bar. Air-flow was guided through polyethylene tubes, ending with a plastic tip, in front of the rat’s whiskers, which were maintained together with a piece of tape.

Stimulation parameters were based on previous rat stimulation fMRI studies and chosen to activate voxels widely in visual and somatosensory pathways, correspondingly (Lu et al., 2016; Van Camp et al., 2006). Both sets of stimulations were initiated manually using the live EEG recording as a guide, either during ictal or interictal state, to pursue equal sampling of both states across each 45-min fMRI session. A delay of at least 20 s was maintained between each stimulation block, to allow hemodynamic responses to settle to baseline between each block. EEG was recorded with a sampling rate of 1024 Hz (Micromed, SD MRI amplifier, Treviso, Italy). EEG recording software received a TTL trigger from the MRI scanner to mark the MRI sequence onset and TTL triggers from the stimulator to mark each stimulation period onset. During the live monitoring, a 50 Hz notch filter, and a low pass filter of 12 Hz facilitated online visual detection of seizures.

Data analysis

Motion analysis during fMRI

Motion occurrences were analyzed by visually inspecting image volumes through each scan and average motions per minute were calculated. Motion correction parameters given by ANTs (see 2.5.3.) were also used and maximum value was taken to estimate maximum displacement of the head from each session. Framewise displacement (FD) (Figure 2E) was calculated as follows. First, the differential of successive motion parameters (x, y, z translation, roll, pitch, yaw rotation) was calculated. Then absolute value was taken from each parameter, and rotational parameters were divided by 5 mm (as estimate of the rat brain radius) to convert degrees to millimeters (Power et al., 2012). Lastly all the parameters were summed together.

EEG analysis

EEG data was converted from native Micromed TRC-file format to mat- and dat-files using Statistical Parametric Mapping (SPM), version 12 (https://www.fil.ion.ucl.ac.uk/spm/software/spm12/, https://www.fil.ion.ucl.ac.uk/spm/doc/spm12_manual.pdf) toolbox ImaGIN (https://github.com/manikbh/ImaGIN2). Data were filtered with a 50 Hz notch and 1-90 Hz Butterworth band pass filters using SPM and ImaGIN. Absence seizures were manually inspected from a filtered signal and spectrogram. Seizures were confirmed as SWDs if they had a typical regular spike and wave pattern with 7-12 Hz frequency range and had at least double the amplitude compared to baseline signal. All other signals were classified as baseline i.e. signal absent of a distinctive 7-12 Hz frequency power but spread within frequencies from 1 to 90 Hz. Two successive seizures were counted as one if there was less than 1 s of baseline signal between them. Seizure was counted as absence seizure only if it lasted at least 2 s, since behavioral deficits are not obvious in shorter seizures (Blumenfeld, 2012). Seizure initiation and ending time points as well as stimulation onsets were marked in seconds and converted to MRI time in volumes.

Temporal registration between EEG and fMRI was performed using TTL triggers delivered by the MRI sequence and registered as the same time as fMRI signals.

fMRI analysis

To prepare the fMRI data for analysis, data were motion corrected, co-registered with anatomical MRI, and spatially smoothed to reduce noise level. The preprocessing steps were performed using advanced normalizing tools (ANTs, http://stnava.github.io/ANTs/) and python based graphical user interphase mri_works (https://montigno.github.io/mri_works/Home/index.html). The first 5 volumes from each fMRI data were removed to allow signal to reach steady state. For motion correction, all volumes were transformed to an average image taken from first 10 volumes by using antsMotionCorr. Normalization was performed indirectly and using the T1-FLASH images, as ZTE images lack anatomical contrast. First, the T1-FLASH images of each animal were corrected from intensity non-uniformity using N4BiasFieldCorrection from ANTs. Next, brain was extracted with MP3 software (Brossard et al. 2020). T1-FLASH images of animals were then co-registered on top of reference T1-FLASH image (taken from a representative animal from our study) with ANTs using rigid, affine (linear) and SYN (non-linear) registrations. Transformation matrices were applied to ZTE images by ANTsApplyTransform so that ZTE images were eventually aligned with the reference T1-FLASH image. Finally, ZTE-images were smoothed with a Gaussian filter (1 mm full width at half maximum).

Earlier version: Post-hoc, stimulation periods were classified into different inputs (Figure 1B, Figure 1–figure supplement 1A), based on the relative position in time of the stimulation period with respect to the absence seizure. Stimulation fully applied during an ictal period and those during an interictal period were used for the ictal-interictal comparison analysis. The effects of the periods when stimulation was ending a seizure (stimulation was considered to end a seizure when there was 0-2 s between stimulation start and seizure end) and the effect of the seizures itself, in absence of stimulation, were also studied. The periods when stimulation ended a seizure are particularly interesting for studying the spatial and temporal aspects explaining shift from ictal, to interictal, i.e., potential non-responsiveness state, to responsiveness state. Intermediate cases, where the seizure started or ended during the stimulation block (Figure 1–figure supplement 1A), were considered as confounds of no-interest in the SPM analysis of fMRI data and the explained variance caused by the confounds were reduced from the main effects of interests. Translational and rotational motion parameters, obtained from a motion correction step, were used also as confounds of no-interest. The variance caused by the confounds of no-interests was reduced from the main effects of interests in the linear model. Stimulations that coincided with a motion above 0.3% of the voxel size were not considered as stimulation inputs. To account for temporal and dispersion variations in the hemodynamic response, stimulation and seizure inputs were convolved with three gamma distribution basis functions (i.e., 3rd order gamma) in SPM (option: basis functions, gamma functions, order: 3). The choice of 3rd order gamma was based on the expectation that time-to-peak and shape of HRFs of seizure could vary across voxels (David et al., 2008).

Revised version: Post-hoc, stimulation periods were classified into different inputs (Figure 1B, Figure 1–figure supplement 1A), based on the relative position in time of the stimulation period with respect to the absence seizure. Stimulation fully applied during an ictal period and those during an interictal period were used for the ictal-interictal comparison analysis. The effects of the periods when stimulation was ending a seizure (stimulation was considered to end a seizure when there was 0-2 s between stimulation start and seizure end) and the effect of the seizures itself, in absence of stimulation, were also studied. The periods when stimulation ended a seizure are particularly interesting for studying the spatial and temporal aspects explaining shift from ictal, to interictal, i.e., potential non-responsiveness state, to responsiveness state. Intermediate cases, where the seizure started or ended during the stimulation block (Figure 1–figure supplement 1A), were considered as confounds of no-interest in the SPM general linear model analysis of fMRI data and the explained variance caused by the confounds were reduced from the main effects of interests. Translational and rotational motion parameters, obtained from a motion correction step, were used also as confounds of no-interest. The variance caused by the confounds of no-interests was reduced from the main effects of interests in the linear model. Stimulations that coincided with a motion above 0.3% of the voxel size were not considered as stimulation inputs. To account for temporal and dispersion variations in the hemodynamic response, stimulation and seizure inputs were convolved with three gamma distribution basis functions (i.e., 3rd order gamma) in SPM (option: basis functions, gamma functions, order: 3). The choice of 3rd order gamma was based on the expectation that time-to-peak and shape of HRFs of seizure could vary across voxels (David et al., 2008).

Earlier version: By convolving each stimulation and seizure inputs with three gamma functions (Figure 1–figure supplement 1B), three regressors from each condition were obtained. and those were used to estimate the regression coefficient values (beta-values) for every voxel in the brain. Next, contrasts were applied to the beta-values as a function of the studied stimulation type. For example, to compare the effect of interictal-ictal stimulations, the beta values corresponding to interictal stimulation were given a contrast of 1, and those corresponding to ictal stimulation were given a contrast of −1. As a result, for each stimulation and seizure condition, three contrast estimates were obtained. From these contrast estimates, statistical F-contrast maps were created and corrected for multiple comparisons by cluster level correction (Figure 3 and 5). F-contrast was used to test differences of any of the three contrast estimates. For difference maps between interictal-ictal stimulations, F-contrast maps were masked with a T-contrast map, where only 1st order basis derivative was used, to reveal brain areas where interictal stimulation provided higher amplitude response than ictal stimulation.

Revised version: By convolving each stimulation and seizure inputs with three gamma functions (Figure 1–figure supplement 1B), three regressors from each condition were obtained. and those were used to estimate the regression coefficient values (beta-values) for every voxel in the brain. Next, contrasts were applied to the beta-values as a function of the studied stimulation type. For example, to compare the effect of interictal-ictal stimulations, the beta values corresponding to interictal stimulation were given a contrast of 1, and those corresponding to ictal stimulation were given a contrast of −1. As a result, for each stimulation and seizure condition, three contrast estimates were obtained. From these contrast estimates, statistical F-contrast maps (look for Friston et al., 1994) were created and corrected for multiple comparisons by cluster level correction (Figure 3 and 5). F-contrast was used to test differences of any of the three contrast estimates. For difference maps between interictal-ictal stimulations, F-contrast maps were masked with a T-contrast map, where only 1st order basis derivative was used, to reveal brain areas where interictal stimulation provided higher amplitude response than ictal stimulation.

Activation F-contrast maps of stimulation responses during interictal and ictal states and difference maps between these two states for the visual (A) and the whisker (B) stimulation experiments.

Parameter estimates of regressors were calculated for every voxel, and contrasts were added to parameter estimates of interictal stimulation, ictal stimulation and to compare interictal versus ictal stimulation in visual (A) and whisker (B) stimulation groups. For statistical significance, F-contrast (p<0.05) maps were created and corrected for multiple comparisons by cluster level correction. AC = auditory cortex, BF = barrel field, mFC = medial frontal cortex, SC = superior colliculus, S1/S2 = primary and secondary somatosensory cortex, Thl = thalamus, VC = visual cortex. a = anterior, p = posterior.

Added in the revised version: HRFs were estimated in ROIs in response to ictal and interictal stimulus (Figure 4). Anatomical ROIs, based on Paxinos atlas (Paxinos and Watson rat brain atlas 7th edition), were drawn on the brain areas where statistical differences were seen in activation maps. Also, HRFs were estimated in response to stimulation that ended a seizure and were compared to the response when stimulation did not end a seizure (Figure 4–figure supplement 1). Three gamma basis functions, for each condition, were multiplied with their corresponding average beta values over a selected ROI, resulting in three functions, and the sum of these functions was taken. The shape and time-to-peak of HRFs during an interictal condition were in line with previous literature (Lambers et al., 2020) thereby confirming the selection of the 3rd order gamma basis function to model HRF. For statistical comparison, extreme beta-values (maximum or minimum) in a selected ROI, between different conditions were compared using a two-tailed t-test (Figure 4).

Hemodynamic response functions to stimulations performed during an interictal and ictal period: visual stimulation (A) and whisker stimulation (B) groups.

HRFs were calculated in selected ROI, belonging to visual or somatosensory area, by multiplying gamma basis functions (Figure 1–figure supplement 1B) with its corresponding average beta value over a ROI and taking a sum of these values. For statistical comparison, extreme beta-values over a ROI were calculated and values between two states were compared with a two-sample t-test. Scatter plot represents mean ± SD and each blue dots corresponds to extreme beta-values observed from individual fMRI sessions. *** = p < 0.001. Grey box illustrates stimulation period.

Removed in the revised version: Activation F-contrast map (A) and hemodynamic response functions to seizure (B).

Removed in the revised version: Parameter estimates of regressors were calculated for every voxel, and contrasts were added to parameter estimates of seizure in absence of stimulation. For statistical significance, F-contrast (p<0.05) maps were created and corrected for multiple comparisons by cluster level correction. HRF was calculated in selected ROI by multiplying gamma basis functions (Figure 1–figure supplement 1B) with their corresponding average beta values over a ROI and taking a sum of these values. For both the maps and HRFs, data from visual and whisker stimulation experiments were pooled together. BF = barrel field, BG = basal ganglia, FC = frontal cortex, MD = mediodorsal thalamic nucleus, VC = visual cortex. a = anterior, p = posterior.

Removed in the revised version: HRFs were estimated in ROIs in response to ictal and interictal stimulus (Figure 4). Anatomical ROIs, based on Paxinos atlas (Paxinos and Watson rat brain atlas 7th edition), were drawn on the brain areas where statistical differences were seen in activation maps. Also, HRFs were estimated in response to stimulation that ended a seizure and were compared to the response when stimulation did not end a seizure (Figure 4–figure supplement 1). Three gamma basis functions, for each condition, were multiplied with their corresponding average beta values over a selected ROI, resulting in three functions, and the sum of these functions was taken. The shape and time-to-peak of HRFs during an interictal condition were in line with previous literature (Lambers et al., 2020) thereby confirming the selection of the 3rd order gamma basis function to model HRF. For statistical comparison, extreme beta-values (maximum or minimum) in a selected ROI, between different conditions were compared using a two-tailed t-test (Figure 4).

Modeling and Simulations

Whole brain simulations of the rat brain during sensory stimulation were performed. Each region of the rat brain was modelled via a recently developed mean-field model of Adaptive Exponential LIF neurons (AdEx) (Di Volo et al., 2019), which describes the activity of a population of neurons made by excitatory and inhibitory cells (for details on the model see Sup. Information and Di Volo et al., 2019). This model has been widely tested and used for the simulation of different brain states (Goldman et al., 2023), neuronal responsiveness (Goldman et al., 2023; Di Volo et al., 2019), and whole-brain dynamics (Goldman et al., 2023). In addition, modeling tools to calculate brain signals (such as LFP and BOLD-fMRI) from this type of mean-fields have been developed (Tesler et al., 2022, 2023). The AdEx mean-field model is capable of producing different neuronal dynamics that can be associated with brain activity during interictal and ictal periods (see results section). The interictal periods were modeled by an asychronous irregular (AI) type of dynamics, while ictal periods were modeled by an oscillatory dynamic which resembles the spike-and-wave discharges (SWD) observed during absence epilepsy. The mean-field model was combined with a connectivity matrix of the rat brain (BAMS rat connectome, Bota et al., 2012) to build a realistic whole-brain simulation of the stimulus propagation during ictal and interictal periods. The simulations were made with the Virtual Brain platform (TVB), which provides a framework to perform large-scale brain simulations (Sanz-Leon et al., 2015).

Results

Fine-Tuning of fMRI recording in awake rats

The spatial SNR was ∼25 whereas the temporal SNR was between 30-60 in the brain (Figure 2). In pilot studies, responses to visual stimulation measured using ZTE and EPI sequences were compared in one animal (Figure 2–figure supplement 1). Average signal change due to stimulation was between 1-1.5% with gradient-echo-EPI sequence while it was ∼0.3% with ZTE. However, activation maps (conducted with Aedes, https://github.com/mjnissi/aedes) demonstrated less susceptibility-induced image distortions in ZTE sequence compared to EPI. Moreover, when studying the effect of fMRI artifacts on EEG signal, more pronounced and higher amplitude gradient switching artifacts were detected when using EPI compared to ZTE sequence.

When comparing peak acoustic noise levels inside the 9.4T Bruker magnet, ZTE (78.7 dB) was 35.8 dB quieter than EPI (114.5 dB) equaling ∼62 times weaker sound pressure (Schomer 1998). Motion during fMRI imaging using ZTE was low (Figure 2b). Mean framewise head translation was only 0.69 ± 0.32 µm and head rotation 0.79 ± 0.55°. Maximum displacement of the rat head was 12.8 ± 11.6 µm, corresponding to 0.03 ± 0.02 voxels, while maximum rotation was 18.1 ± 12.9°. Motion occurrences between the current study (0.43 ± 0.45 motions/min) and the previous study (1.0 ± 0.20 motions/min), performed with a similar rat restraint holder but an EPI sequence (Stenroos et al., 2018), demonstrated the advantage of using the quieter ZTE sequence compared to the louder EPI sequence. ZTE motion levels were similar to that obtained in a previous study using a quiet MB-SWIFT sequence (0.48 ± 0.23 motions/min, Paasonen et al., 2020).

In our study, use of 3 out of 11 animals (27%) had to be discontinued due noisy EEG signal, most likely caused by partly detached implant. One animal (1%) was excluded due to a lost implant. From MRI-measurements, 6 out of 28 sessions (21%) needed to be excluded from analyses due to an excessive movement that prevented reliable fMRI analysis. Moreover, 4 sessions (14%) were excluded due to lack of seizures, and 1 session (4%) was excluded due to technical failure in providing air-puffs to whiskers. Respiration of awake animals remained stable during the measurements. Average respiration frequency across the 45-min scans was 2.2 ± 0.4 Hz in visual stimulation group and 2.0 ± 0.3 Hz in whisker stimulation group and no sudden changes of respiration due to stimulations were noted.

Response to sensory stimuli

Earlier version: Stimulations were manually initiated during ictal and interictal periods, but some stimulation blocks co-occurred unintentionally in-between ictal and interictal periods. Table 1 shows the amount of each stimulation type and the amount and duration of seizure periods.

Revised version: Stimulations were manually initiated during ictal and interictal periods, but some stimulation blocks co-occurred unintentionally in-between ictal and interictal periods. Table 1 shows the amount of each stimulation type and the amount and duration of seizure periods. Regarding statistical power, considering a risk alpha of 0.05, a power of 0.8, matched pairs (seizure/control), we can detect an effect size of 0.37 with 4 animals, considering repeated measurements (4 sessions/animal x 11 seizure/control pairs per session).

Earlier version: Statistical activation maps (Figure 3) in response to stimulus were created in interictal and ictal periods and compared between these two states. During interictal condition in the visual stimulation group, statistical responses (p<0.05, cluster-level corrected) were most notably seen in the visual cortex, the superior colliculus, the thalamus (including the lateral geniculate nucleus) and the frontal cortex (including the prelimbic, cingulate and secondary motor cortices). However, during a seizure condition, responses in the visual cortex were less pronounced while responses in the superior colliculus remained stable. In addition, responses were observed in the somatosensory cortex during a seizure state. There were more activated voxels during interictal state compared to ictal state (136% more). Comparing the statistical responses between interictal and ictal states revealed significant changes (p<0.05, cluster-level corrected) in the visual, somatosensory, and medial frontal cortices: these regions were less activated ictal than interictal state (see also Figure 4).

Revised version: Statistical activation maps (Figure 3) in response to stimulus were created in interictal and ictal periods and compared between these two states. During interictal condition in the visual stimulation group, statistical responses (p<0.05, cluster-level corrected) were most notably seen in the visual cortex, the superior colliculus, the thalamus (including the lateral geniculate nucleus) and the frontal cortex (including the prelimbic, cingulate and secondary motor cortices). However, during a seizure condition, responses in the visual cortex were less pronounced while responses in the superior colliculus remained stable. There were more voxels with significant changes of activity during interictal state compared to ictal state (136% more). Comparing the statistical responses between interictal and ictal states revealed significant changes (p<0.05, cluster-level corrected) in the visual, somatosensory, and medial frontal cortices. In the ictal state, these regions were showed significant hemodynamic decreases when comparing to interictal state, and these polarity changes can be seen the hemodynamic response functions (Figure 4).”

Earlier version: During the interictal condition in the whisker stimulation group, responses (p<0.05, cluster-level corrected) were most notably seen in the somatosensory cortex, including the barrel field, the auditory cortex, the thalamus, including the ventral posteromedial nucleus, and the medial frontal cortex. During a seizure, responses due to stimulations were less pronounced in the thalamus and the frontal cortex and still present in the somatosensory cortex. More voxels were activated during an interictal state compared to a seizure state (179% more). When comparing statistical responses between both states, significant changes (p<0.05, cluster-level corrected) were noticed in the somatosensory, auditory and frontal cortices: these regions were less activated in ictal than in interictal state (see also Figure 4).

Revised version: During the interictal condition in the whisker stimulation group, responses (p<0.05, cluster-level corrected) were most notably seen in the somatosensory cortex, including the barrel field, the auditory cortex, the thalamus, including the ventral posteromedial nucleus, and the medial frontal cortex. During a seizure, responses due to stimulations were less pronounced in the thalamus and the frontal cortex and still present in the somatosensory cortex. There were more voxels with significant changes of activity during interictal state compared to ictal state (179% more). When comparing statistical responses between both states, significant changes (p<0.05, cluster-level corrected) were noticed in the somatosensory, auditory and frontal cortices: these regions showed significant hemodynamic decreases in ictal state compared to interictal state (see also Figure 4).

Analyses of brain hemodynamic response functions to sensory stimulations during ictal and interictal brain state

In the visual stimulation group, extreme beta-value of the response in the visual cortex ROI was 2.8 ± 1.7 during the baseline period, and −6.0 ± 2.0 during the seizure period, with a significant difference between the two conditions (p < 0.001). Additionally, the response amplitude was higher when the stimulation ended a seizure compared to when it did not (8.1 ± 7.0 vs −6.0 ± 2.0, p < 0.001). In the whisker stimulation group, the extreme beta-value of the response in the barrel cortex was 4.1 ± 1.9 during the baseline period, and −9.0 ± 1.9 during a seizure, also with a significant difference between the two conditions (p < 0.001). In this group, stimulation responses were also higher when the stimulation ended a seizure compared to when it did not (4.8 ± 2.9 vs −9.0 ± 1.9, p < 0.001). HRFs amplitudes were both negatively and positively signed during the ictal state, depending on the brain region.

When analyzing the effect of a seizure itself, the sessions from the visual and whisker stimulation experiments were pooled together. Significant changes (p<0.05, cluster-level corrected) were observed in various cortical areas, including the frontal, parietal and occipital regions (Figure 5A) totaling 895 voxels. Response amplitudes were predominantly negatively signed in the cortical regions, while they were positively signed in deeper brain regions, such as thalamus and basal ganglia (Figure 5B).

Added in the revised version: Activation F-contrast map (A) and hemodynamic response functions to seizure (B).

Added in the revised version: Parameter estimates of regressors were calculated for every voxel, and contrasts were added to parameter estimates of seizure in absence of stimulation. For statistical significance, F-contrast (p<0.05) maps were created and corrected for multiple comparisons by cluster level correction. HRF was calculated in selected ROI by multiplying gamma basis functions (Figure 1–figure supplement 1B) with their corresponding average beta values over a ROI and taking a sum of these values. For both the maps and HRFs, data from visual and whisker stimulation experiments were pooled together. BF = barrel field, BG = basal ganglia, FC = frontal cortex, MD = mediodorsal thalamic nucleus, VC = visual cortex. a = anterior, p = posterior.

Simulations of sensory stimulation during ictal and inter-ictal periods

Figure 6 shows the results of the simulations. First, neuronal activity is obtained from the AdEx mean-field for the interictal (Figure 6A) and ictal periods (Figure 6B). For a better description, the local field potential (LFP) and the membrane potential, calculated from the mean-field model during the SWD type of dynamics are shown in Figure 6C. The LFP was computed using a recently developed kernel method (Tesler et al. 2022). We see that the model can capture a SWD type of pattern, similar to the one observed experimentally in electrophysiological measurements, which is correlated with periods of hyper-polarization in the membrane potential (Figure 6, Figure 6–figure supplement 2). In the model the hyper-polarization is driven by a strong adaptation current and the SWD dynamics is suppressed by reducing the strength of this current. Thus, the switch between AI and SWD dynamics in our model is given by varying the strength of the adaptation current.

Simulation of sensory stimulation during ictal and interictal periods.

A-B: asynchronous irregular (AI) and spike-and-wave discharge type of dynamics obtained from the mean-field model, representing interictal and ictal periods respectively. The change between the two dynamics is given by the strength of the adaptation current in the AdEx mean-field model. C: LFP and membrane potential obtained from the mean-field model. The model can capture the SWD pattern observed experimentally in LFP measurements which is correlated with periods of hyper-polarization in the membrane potential. D-E: time-series and statistical maps of the simulated sensory stimulus in the whole brain simulations of the rat, showing the results of a stimulation of the primary visual cortex during ictal and inter-ictal periods. The onset and duration of the stimulus is indicated by the dashed vertical line and horizontal line at the top of the time-series. The statistical maps are built from a 2D representation of the 496 brain regions of the BAMS rat connectome described in the method section 2.6.

The time series and statistical maps of the whole brain simulations in response to stimulus, performed with the TVB platform (see Methods), are described in Figures 6D and E. The statistical maps for the simulations are calculated directly from the neuronal activity (firing rates). The stimulation of a specific region is simulated as an increase in the excitatory input to the specific node. In particular we use a periodic square function for representing the stimulus (see panel A in Figure 6–figure supplement 1). For the results presented here, the stimulus was simulated in the primary visual cortex (indicated by the red-circled pixel). As we can see from the statistical maps, the propagation of the stimulus is drastically restrained during ictal periods in comparison with interictal periods. During the ictal periods the dynamics of the system is dominated by the highly synchronous and regular SWD oscillation, and the effect of the stimulus does not alter significantly the ongoing dynamics. On the contrary, during the interictal periods, where the system exhibits an asynchronous irregular dynamic, the effect of stimulus generates a large variation of the ongoing dynamics in the regions linked to the stimulated area (see also Figure 6–figure supplement 1). This agrees with the reduced responsiveness observed in the fMRI results for stimulation during ictal periods.

Discussion

The investigation of how sensory stimulations are handled by the brain in case of absence epilepsy is essential to develop better care for individuals with this pathology and associated comorbidities. Thanks to specific fMRI and EEG recordings obtained in awake rats, we have uncovered significant differences in brain activity and activation of specific structures during visual and whisker stimulations, depending on whether they occur during a seizure period or not.

Zero echo time fMRI of non-curarized awake rats enabled absence seizure detection

Recording absence seizures in rats can be challenging because they only occur when the animals are in a calm and stress-free awake state. Despite the disturbance caused by MRI scanning noise and the need for rat restraint, we were able to successfully detect absence seizures in awake GAERS during training periods and fMRI sessions. We found that the ZTE sequence, which produced considerably lower sound pressures compared to traditional EPI, was the optimal MRI sequence for our purpose. Based on our experience from current and previous experiments, we believe that acoustic noise level is the most significant stress factor for awake rats because it can cause increased motion and thereby potential confounds in task-based and resting-state fMRI studies. ZTE-based fMRI may provide better localization of activated sites because it is sensitive to change in blood flow (Luh et al., 2000; Restom et al., 2007), and it has also been suggested that ZTE-fMRI is 67% more sensitive than standard BOLD EPI due to its ability to detect increases in tissue oxygenation which shortens the T1-relaxation rate of spins in a pseudo steady-state (MacKinnon et al. ISMRM 2021).

Whole brain responsiveness during an interictal condition

In the visual stimulation group, responses were most pronounced in the visual cortex, the superior colliculus and the thalamus, including the lateral geniculate nucleus, which are all part of the rat visual pathways (Sefton et al., 2015). As binocular stimulation was used, responses in both hemispheres were observed. Interestingly, we also detected activation in the frontal cortex, including the prelimbic, cingulate and secondary motor cortices. The involvement of frontal cortical areas can be explained by the fact that medial prefrontal cortex, which includes the prelimbic cortex, is known to be a central area for rat’s attention (Williams et al., 1999) and that lesion in the medial prefrontal cortex impaired the accuracy of detecting brief flashes of light (Muir et al., 1996).

In the whisker stimulation group, responses were mostly seen in the somatosensory cortex, including the barrel field, and in the ventral posteromedial thalamus, which are part of the rat whisker system (Adibi, 2019). This finding suggests that the somatosensory cortex, the initiating zone of the seizures, is functional during the interictal state in GAERS. Additionally, we detected activation in the frontal cortex, which could be due to increased attention to external cues. Surprisingly, positive activation was also observed in part of the auditory cortex. This activation could be explained by the rats’ ability to differentiate the sound coming from air puffs from the ongoing MRI-noises, causing increased auditory activation during stimulation.

Decreased cortical activity during a seizure per se

Experimental and simulation results (Figures 1 and 6) illustrated typical synchronous and regular SWD patterns with a spike component followed by a longer-lasting wave or silence component. The silence component of SWD is thought to be caused by neuronal deactivation or increased inhibitory activity (Fisher & Prince, 1977; Gloor et al. 1978; Inoue et al., 1993; Kostopoulos et al., 1982; McCafferty et al., 2023), potentially resulting in overall decreased neuronal activity in the thalamocortical and cortical neurons. Our findings of negative HRFs in the cortical regions (Figure 5) are in line with this hypothesis, suggesting reduced neuronal activity. Supporting evidence was found in a recent neuroimaging study conducted with GAERS showing that during an absence seizure, a larger group of neural populations in the frontoparietal cortex had a decreased firing rate, while a smaller portion of neuronal populations had an increased firing rate, leading to reduced fMRI cortical activity (McCafferty et al., 2023). Moreover, other studies utilizing optical flowmetry in GAERS have noticed decreased cerebral blood flow in cortical capillaries during an absence seizure, with preserved PO2, PCO2, and arterial blood pressure (Neblig et al., 1996) suggesting decreased cortical brain activity. However, the mechanism for the observed subcortical fMRI increases in this study and a previous study (McCafferty et al., 2023) illustrating the disagreement between electrophysiological and fMRI signals, are yet to explained and demand further studies.

Decreased responsiveness during an absence seizure

Based on fMRI results, we noticed reduced activation in the cortex during ictal state, along with limited propagation of activity compared to the interictal period (Figure 4), which could indicate decreased responsiveness and information processing during external stimulation. Simulation results using a mean-field model also illustrated no observable change in spiking rate together with restricted propagation of neuronal activity in response to visual stimulation (Figure 6 D and E, Figure 6–figure supplement 1), confirming fMRI finding. Previous electrophysiological experimental studies using whisker stimulation setups on GAERS have demonstrated that event-related potentials are modulated but still present during an absence seizure compared to seizure-free periods (Chipaux et al., 2013; Studer et al., 2019). Therefore, our findings do not necessarily correspond to a decreased amplitude of event-related-potential per se but could rather indicate decreased sum of excitatory-inhibitory neuronal state when stimulating during SWD.

During a SWD, the overall neuronal activity in several cortical areas is globally decreased (McCafferty et al., 2023) potentially due to the wave component of SWD with increased thalamocortical inhibition (McCafferty et al., 2018). In that state, excitation caused by external stimulation can cause relative increases of activity in responsive brain areas. However, if excitation in these neuronal populations is relatively lower compared to the neuronal suppression caused by SWD, the total sum of activity is negative compared to baseline. As fMRI is believed to measure the sum of neuronal activation (Logothetis et al., 2001), this appears as reduced fMRI amplitude, which was also apparent in measured HRFs in cortical regions (Figure 4). Thus, reduced fMRI responsiveness implies that SWD is a dominant brain feature even under stimulation condition, and that particularly the ‘wave’ phase of the neuronal oscillatory pattern of SWD can prevent responsiveness in GAERS during these conditions.

As previously discussed by Chipaux and collaborators (Chipaux et al., 2013), the lack of conscious perception to paroxysmal stimulation event could be caused by the concomitant appearance of SWD or SWD complex that is shortly (<500 ms in humans) following a stimulation. In our study, as SWD complexes co-occurred with stimulation events, this could hinder the potential conscious perception of visual and somatosensory stimulation. Interestingly, when stimulation events were applied during a seizure but were not followed by SWD complexes (i.e. stimulation ended a seizure), fMRI response amplitudes were higher than when stimulation did not end a seizure. This could mean that the stimulation response exceeded a threshold to shift the brain back to a responsive or conscious state, which was apparent in the fMRI signal as well. The reason for the change in brain state in these stimulation cases is unclear, as stimulation strength remained the same throughout the experiment. However, there is a potential for future experiments and therapeutic interventions to study further how different stimulation schemes affect neuronal populations and cause a shift from ictal to non-ictal state.

Conclusion

In this study, we were able to detect spontaneous absence seizures inside the magnet using awake, non-curarized, rats and a quiet ZTE imaging sequence, and thereby able to study whole brain sensory perception in ictal and interictal brain states. Rats were stimulated with blue led light and with somatosensory whisker air-puffs during both interictal and ictal states. We found that cortical activation was spatially wider and stronger during interictal stimulation compared to ictal stimulation. Moreover, the detected decreases in the cortical HRF when sensory stimulation was applied during spike-and-wave discharges, could play a role in decreased sensory perception. Further studies are required to evaluate whether this HRF change is a cause or a consequence of the reduced neuronal response. Altogether, these stimulation-based fMRI experiments, combined with the brain-wide simulation results, suggest that rats with absence epilepsy have hindered and restricted sensory processing during an ongoing seizure, although not fully abolished. These findings may contribute to our understanding of cognitive impairments observed in patients suffering from absence epilepsy and the severe neuropsychiatric comorbidities observed in children. Additionally, the results of this study provide insight into the information processing during SWD and may aid in the development of future therapeutic approaches for absence epilepsy.

Supporting information